Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 3: Counting Stacks and Queues



Submission Deadline: Wednesday May 17, 11pm. 
Submission link: TBD

## Stacks, Queues, and Their Counting Versions

A stack is a data structure with two operations: push, and pop.  Picture it as a pile of dishes sitting on a counter.  A push operation places a dish on top of the pile.  A pop operation returns the dish on top of the pile, or None if the pile is empty, that is, contains no dishes.  A "dish" can be any Python object. 

A queue is a data structure with two operations: put, and get.  Imagine it as a stack of books horizontally on a shelf.  A put operation adds the book to the left end of the books on the shelf; a get operation gets the book from the right end of the shelf.  

Thus, the difference between a stack and a queue is that the stack is FILO (First In, Last Out), whereas the queue is FIFO (First In, First Out).  Elements in a stack are retrieved newest first. 
Elements in a queue are retrieved in the order they were put in, oldest first.

We will implement here these data structures, with a small twist: we will also introduce _counting_ versions of them, which avoid keeping multiple identical copies of objects in a row. 



Here is the implementation of the Counting Queue we defined in class. 

```python
class CountingQueue(object):

    def __init__(self):
        self.queue = []

    def __repr__(self):
        return repr(self.queue)

    def add(self, x, count=1):
        # If the element is the same as the last element, we simply
        # increment the count.  This assumes we can test equality of
        # elements.
        if len(self.queue) > 0:
            xx, cc = self.queue[-1]
            if xx == x:
                self.queue[-1] = (xx, cc + count)
            else:
                self.queue.append((x, count))
        else:
            self.queue = [(x, count)]

    def get(self):
        if len(self.queue) == 0:
            return None
        x, c = self.queue[0]
        if c == 1:
            self.queue.pop(0)
            return x
        else:
            self.queue[0] = (x, c - 1)
            return x

    def isempty(self):
        # Since the count of an element is never 0, we can just check
        # whether the queue is empty.
        return len(self.queue) == 0


```

Let's put this to the same test as before, printing the queue contents at each step to see what is going on.

```python
q = CountingQueue()
q.add('a')
print(q)
q.add('b', count=5)
print(q)
q.add('c', count=2)
print(q)
while not q.isempty():
    print(q.get())
    print(q)

```

Output:
```
[('a', 1)]
[('a', 1), ('b', 5)]
[('a', 1), ('b', 5), ('c', 2)]
a
[('b', 5), ('c', 2)]
b
[('b', 4), ('c', 2)]
b
[('b', 3), ('c', 2)]
b
[('b', 2), ('c', 2)]
b
[('b', 1), ('c', 2)]
b
[('c', 2)]
c
[('c', 1)]
c
[]
```

```python
#@title Testing helper

def check_equal(x, y, msg=None):
    if x == y:
        if msg is not None:
            print(msg, ": Success")
    else:
        if msg is None:
            print("Error:")
        else:
            print("Error in", msg, ":")
        print("    Your answer was:", x)
        print("    Correct answer: ", y)
    assert x == y, "%r and %r are different" % (x, y)

```

## Question 1: Implement the `__len__` Method

If you want to take the length of an object, as in 

    len(someobject)

then `someobject` must have a `__len__` method, which should return the length. 
Implement the `__len__` method for counting queues, so that it returns the queue length (the number of elements in it, counting repeated elements as distinct). 

<!-- grade_id: cell-c441c9c06ba06163, points: 0, grade: false, solution: true -->
```python
### Exercise: implement `__len__` for a counting queue

def countingqueue_len(self):
    """Returns the number of elements in the queue."""
    # YOUR CODE HERE

# This is a way to add a method to a class once the class
# has already been defined.
CountingQueue.__len__ = countingqueue_len

```

<!-- grade_id: cell-3cba967a4571f121, points: 10, grade: true, solution: false -->
```python
### Tests for `__len__`

q = CountingQueue()
for i in range(10):
    q.add('a')
q.add('b')
for i in range(3):
    q.add('c', count=2)
check_equal(len(q), 17)
check_equal(len(q), 17)
```

Output:
```
TypeError                                 Traceback (most recent call last)
/var/folders/tq/0p82818x4jd9plq_7b32kjl00000gn/T/ipykernel_58100/2629145710.py in <module>
      7 for i in range(3):
      8     q.add('c', count=2)
----> 9 check_equal(len(q), 17)
     10 check_equal(len(q), 17)

TypeError: 'NoneType' object cannot be interpreted as an integer
```

## Question 2: implement the `__iter__` Method

We would like to be able to have a way of iterating over elements in our counting queue.  
This can be used, for instance, to print them, or to process the elements in some way. 

Precisely, we would like to have a way of writing, for a counting queue `q`:

    for el in q:
        print el

and we would like this to print all queue elements, in order. 

The way to achieve this is to define an `__iter__` method that acts as a generator for the elements. 

<!-- grade_id: cell-1d6ded581486b321, points: 0, grade: false, solution: true -->
```python
### Exercise: Write an iterator for CountingQueue

# Note: it can be done elegantly in 3 lines of code.

def countingqueue_iter_elements(self):
    """Iterates through all the elements of the queue,
    without removing them."""
    # YOUR CODE HERE

CountingQueue.__iter__ = countingqueue_iter_elements

```

<!-- grade_id: cell-9394eb3e713b0c74, points: 10, grade: true, solution: false -->
```python
### Tests for `CountingQueue.__iter__`

q = CountingQueue()
for i in range(10):
    q.add('a')
q.add('b')
for i in range(3):
    q.add('c', count=2)
l1 = [x for x in q]
l2 = []
l3 = [x for x in q]
assert len(l3) > 0
while not q.isempty():
    l2.append(q.get())
check_equal(l1, l2)


```

Output:
```
[]
['a', 'a', 'a', 'a', 'a', 'a', 'a', 'a', 'a', 'a', 'b', 'c', 'c', 'c', 'c', 'c', 'c']
```